# RandomLogistic
In this assigment you will implement RandomLogistic Regression. 
- First you will read the image using the given instructions.
- Then you will flatten and normalize the image.
- After you implement model compute function will be used to get the result. 

## **IMPORTANT WARNINGS**
- **DO NOT USE** any additional libraries other than the standard Python libraries and the ones in the requirements.txt
- **DO NOT DELETE** any cells even if it they are empty
- While implementing the functions you should only change the cells with the statement #YOUR CODE HERE \
- If you change other than those blocks your notebook might not be graded properly

```python
import numpy as np
from PIL import Image
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

## Image reading
Read the image using PIL library and convert the image to a numpy array using 'np.asarray' function

<!-- grade_id: cell-a6681d0d24df745a, points: 0, grade: false, solution: true -->
```python
def read_image(img_path):
    return  np.asarray(Image.open(img_path), dtype=None, order=None)
```

```python
img = read_image('inzva.jpeg')
```

<!-- grade_id: cell-27634dfa545661ce, points: 10, grade: true, solution: false -->
```python
assert type(img) == np.ndarray
assert img.shape == (225, 225, 3)
```

## Flatten the Image
The given image has the shape (1000, 1000, 3). We want you to reshape the image into (3000000,)

<!-- grade_id: cell-52062b23598241d2, points: 0, grade: false, solution: true -->
```python
def flatten_img(img):
    # YOUR CODE HERE
    return img.flatten()
```

```python
img_flatten = flatten_img(img)
```

<!-- grade_id: cell-7b2b4622d95c08a9, points: 10, grade: true, solution: false -->
```python
assert img_flatten.shape == (151875,) or img_flatten.shape == (151875,1)
```

## Normalize the Image
You should implement the following formula to normalize the image:
$$ x = \frac{x - min(x)}{max(x)-min(x)} $$

<!-- grade_id: cell-9b4666d8308a183b, points: 0, grade: false, solution: true -->
```python
def normalize_img(img):
    # YOUR CODE HERE
    return (img - np.min(img)) / (np.max(img) - np.min(img))
```

```python
img_normalized = normalize_img(img_flatten)
print(img_normalized.shape)
```

Output:
```
(151875,)
```

<!-- grade_id: cell-61a350c6477164bb, points: 10, grade: true, solution: false -->
```python
assert np.max(img_normalized) == 1.0 and np.min(img_normalized) == 0.0
```

## Random Logistic
Set weights using 'np.random.randn' and set bias to zero. Be careful with the dimensions. Then implement the sigmoid function ($\sigma $) using the following formula:
$$\sigma(x) = \frac{1}{ 1 + \exp(-x)} $$
Then implement the forward pass by:
$$y = \sigma((W * x) + b)$$

<!-- grade_id: cell-7ef2b26abd49de09, points: 0, grade: false, solution: true -->
```python
class RandomLogistic:
    def __init__(self, input_shape):
        np.random.seed(2024)
        ## Define your weights and bias below.
        # YOUR CODE HERE
        self.W = np.random.randn(*input_shape)
        self.W0 = np.zeros(1)
         

    def sigmoid(self, x):
        return 1/(1+np.exp(-x))

    def compute(self, x):
        return self.sigmoid((self.W * x) + self.W0)
```

```python
random = RandomLogistic(input_shape=img_flatten.shape)
y = random.compute(img_normalized)
z = np.load('result.npy')

```

```python
print(y)
print(y.shape)

```

Output:
```
[0.50817597 0.50361439 0.49901207 ... 0.50156481 0.5003981  0.51173913]
(151875,)
```

```python
print(z)
```

Output:
```
[0.50817597 0.50361439 0.49901207 ... 0.50156481 0.5003981  0.51173913]
```

<!-- grade_id: cell-3a57a78b3ebce3d2, points: 20, grade: true, solution: false -->
```python
assert np.sum(np.abs(y-z)) < 1e-9
```

# Edit Distance

- Complete the function below to find the edit distance between two strings.
- Edit distance between two words is defined as the minimum number of operations needed to transform one word into another, in which the only operations allowed are inserting, deleting or replacing a character. 
(insertion: inzva → inz**i**va, deletion: inz**v**a → inza, replacement: inzva → in**o**va)

**For this question, use Python’s standard libraries only.**\
Make sure that your algorithm works correctly as there are also hidden test cases that are not shown here.

<!-- grade_id: cell-3f5cca84206b51d5, points: 0, grade: false, solution: true -->
```python
def edit_distance(word1, word2):
    l1 = len(word1)
    l2 = len(word2)
    dp = [[0 for _ in range(l2 + 1)] for _ in range(l1 + 1)]
    for i in range(1, l1 + 1):
        dp[i][0] = i
    for j in range(1, l2 + 1):
        dp[0][j] = j
    for i in range(1, l1 + 1):
        for j in range(1, l2 + 1):
            if word1[i - 1] == word2[j - 1]:
                dp[i][j] = dp[i - 1][j - 1]
            else:
                dp[i][j] = min(dp[i - 1][j - 1], dp[i - 1][j], dp[i][j - 1]) + 1
    
    return dp[l1][l2]

# print(edit_distance("inzva", "dlsg"))
```

Output:
```
5
```

<!-- grade_id: cell-e68c3f781e28fcf7, points: 6, grade: true, solution: false -->
```python
assert edit_distance("inzva", "inziva") == 1
assert edit_distance("inzva", "inova") == 1
assert edit_distance("inzva", "dlsg") == 5
```

<!-- grade_id: cell-68a823ce760d1540, points: 15, grade: true, solution: false -->
```python
### DO NOT DELETE THIS CELL
```

<!-- grade_id: cell-f2b3f99ef05313d8, points: 29, grade: true, solution: false -->
```python
### DO NOT DELETE THIS CELL
```